```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 4 – Advanced Table Methods

## Data 6, Summer 2025

In this lab, we will be talking all about **tables**. We use tables to store all sorts of data form sports statistics to population information. If there's data you have ever been curious about, it is very likely that the internet has a table somewhere with that data!

Tables are integral to the foundation of Data Science, and in this lab we will go over how to query a table. Querying a table is basically asking information about the table. Some examples of common queries (in English, not code):

* How many data points are there?
* Which data points have a specific characteristic?
* What is the attribute of a specific data point?
* And many more!


Then we will transition to some more complex table methods! The `apply`, `group`, `pivot`, and `join` methods all allow us to perform different queries on our familiar tables. Understanding not only *how* each method works, but also *why* and *when* to use them are the key takeaways from this lab; by the end of it, we will be able to query tables in some pretty cool ways!

These new methods allow us to do different operations than before. As such, it is becoming more and more important to remember how each method sits on our Data Science toolbelt. We should think of each new method as a **tool that serves a specific purpose**. Your job as a data scientist is not only to understand what each tool does, but when each tool is applicable in new situations!

```python
from datascience import *
import numpy as np

import matplotlib.pyplot as plt
plt.style.use("ggplot")
%matplotlib inline
```

<div class="alert alert-warning">
The data that we'll be using in this lab contains data on <a href=https://ir.ucr.edu/stats/admission/undergraduate>University of California, Riverside's undergraduate admissions.</a> We'll be looking at this data today as we are interested in exploring the context and ethics behind education-related data (UC Riverside specifically was chosen as it was the UC that provided the most publicly accessible method of downloading the data). In this case, we are looking at the demographics of undergraduate students. This leads us to many important questions to consider: why does UC Riverside make this data publicly available? Why does it choose this format and these categories to represent students? Who or what is affected by the conclusions we make from analyzing this data? Think about these questions as you work through the lab.
</div>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part 1: Basic Table Functions

<div class="alert alert-warning">
In the first part of this lab, we'll be learning about some of the various basic table functions that we can utilize in the <code>datascience</code> library in order to answer questions about student demographics at UC Riverside. In this part, we'll be focusing on the demographic categories of <code>Gender</code> and <code>IPEDS Race/Ethnicity</code> to explore what the breakdown of students at UC Riverside looks like, as well as getting a better understanding of how and why UC Riverside chose to represent their data publicly in this format.
</div>

### Loading a Table

Recall in Lab 2, we introduced the `Table.read_table` method, which takes a *file path* and constructs a `Table` with the information from that file. Let's see how this works using the file `"riverside.csv"`, which contains information about UC Riverside's admissions.

*Note*: If you want to check where the `riverside.csv` file is, you can look in your DataHub directory by clicking `File` > `Open...` in the top left.

```python
ucr = Table.read_table("riverside.csv")
ucr.show(5)
```

<!-- BEGIN QUESTION -->

### Question 1.1 (Discussion)
As we discussed in the previous lab, what is the unit of analysis of our dataset? Another common way of asking this when considering our data is: what is the granularity of our data?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 1.2
Now that we have considered what each row in the data represents, let's take a look at how many rows we have overall. This gives us a better idea of approximately how many years and stratifications of students per year we have in our data. Assign the variable `ucr_rows` to the number of rows in `ucr`. 

You should not write an integer, but instead use one of the table attributes we have talked about so far to **calculate** the number of rows.

*Hint:* Stuck? Remember, you can reference all of the Table tools in `datascience` by looking at the Data 6 Python Reference sheet [here](http://data6.org/su25/reference). 

```python
ucr_rows = ...
ucr_rows
```

```python
grader.check("q1_2")
```

<div class="alert alert-warning">
To get our analysis started, let's take a look at the <code>drop</code> and <code>relabeled</code> methods to see how we can clean up the appearance of the columns and tables to make viewing and working with the data easier.
</div>

---
### Excluding columns: `drop`

We now have information about UC Riverside's admissions in a `datascience` Table. It is common in data analysis for our tables to have columns that may not be relevant to what our exploration is focused on. One way of cleaning up our table to remove these unnecessary columns is to use the `drop` method!

Along with having columns that may not be relevant to our exploration, it is very common for our datasets to have many rows or columns with `NaN` values (`nan` stands for "Not A Number", and it is typically used to denote a missing value). **However**, it is not a good idea to blindly drop columns with several NaN values from a table. Think back to what you saw with the missing values in Lab 2. What information would have been lost if we just dropped all missing values?

Our table below fortunately doesn't have any missing values, so we'll be using the `drop` method to get rid of columns that aren't relevant in our analysis. This will make it easier to focus on the columns that we do care about. Below, let's use the `drop` method to remove a column from our table, such as the `"New/Continuing Status"` column:

```python
ucr = ucr.drop("New/Continuing Status")
ucr.show(5)
```

*Note:* By dropping the `"New/Continuing Status"` column above, we no longer have information on whether the students in the stratification are new or continuing, so there may be some rows that appear to be duplicates. This is one of the downsides of dropping columns we aren't focusing on -- we may lose some data and context in the process! However, for our exploration's purposes, this information isn't that important, so we'll continue without it.

### Question 1.3 (Multiple Choice)
If the `"Full-Time Status"` column is dropped using the `drop` method, what would `ucr_dropped` look like? You can answer the question by assigning `q1_3` to `make_array(...)` where `...` is the choice of your answer (eg. `make_array(3)`).

1. It contains only rows with NaN values.
2. It contains only the columns before the `"Full-Time Status"` column.
3. It contains all columns except the `"Full-Time Status"` column.
4. It contains the same columns as the original table, but all NaN values in the `"Full-Time Status"` column are dropped.

```python
q1_3 = ...
```

```python
grader.check("q1_3")
```

---
### Changing column labels: `relabeled`

Sometimes the original column names of the data are long or not descriptive enough. We can rename column labels using the `relabeled` method. With this function, you are able to:
1. Relabel a *single column*
2. Relabel *several columns* at once

To change the names of multiple columns, we pass in an array of the old names and an array of the new names as the 2 inputs to `relabeled`.

*Note*: You may see another method called `relabel` in the `datascience` documentation. Please avoid using this, as it can change your data when you may not want to.

### Question 1.4
If we're accessing the `"IPEDS Race/Ethnicity"` column a lot, it could be cumbersome to type the long column name multiple times. Let's change the column label instead to `"Race/Ethnicity"`.

*Hint*: We've provided skeleton code for you to use.

```python
ucr_relabeled = ucr.relabeled(..., ...)
ucr_relabeled.show(5)
```

```python
grader.check("q1_4")
```

<div class="alert alert-warning">
In the next two sections of this part, we'll be looking at the <code>column</code> and <code>select</code> methods to discover ways of focusing on certain subsections of columns in our data.
</div>

---
### Querying the table: `column`

Looking at the first five rows of `ucr_relabeled` above, we can see some of the possible values for each of the columns. For example, in the `"Race/Ethnicity"` column, we see that there are values such as "Asian", "Black/African American", etc. In order to get a better idea of the individual values that exist within a column of our dataset, we can query the table using the `column` method. Let's use this method on `"Race/Ethnicity"`.

```python
ethnicity_list = ucr_relabeled.column("Race/Ethnicity")
ethnicity_list
```

One downside of doing it this way is that the list would be long and repetitive, since it will include the race/ethnicity for every row, resulting in repeated values. To work around this, we can use the `np.unique` method to tell us all the race/ethnicity values only once as they appear:

```python
np.unique(ethnicity_list)
```

<div class="alert alert-warning">
This gives us a much clearer picture of what the possible representations of race/ethnicity are in this dataset. Looking at the values above, here's something to consider: do you find any issues with how the individuals can be categorized? Are there any labels that could be misleading or misrepresentative? Are there any groups of people who might not be included by this breakdown? <br><br>

We'll consider this idea of representation of the data as we continue through the lab.
</div>

---
### Picking columns: `select`

In many cases during data analysis, there will be several columns that we aren't interested in, and only a few that we are. Instead of dropping several columns as we learned about earlier, we can use the `select` method to grab only the columns we want. 

### Question 1.5
For this example, we only want to keep the `"Year"`, `"Level"`, `"Full-Time Status"`, `"Race/Ethnicity"`, and `"Fall Headcount"`  columns. Fill in the following code so that the `ucr_relabeled` table has only the relevant columns.

```python
ucr_selected = ...
ucr_selected
```

```python
grader.check("q1_5")
```

### Question 1.6 (Multiple Choice)
What is the purpose of using the `select` method in this context? You can answer the question by assigning `q1_6` to `make_array(...)` where `...` is the choice of your answer (eg. `make_array(3)`).

1. To drop specific columns from the table selectively.
2. To filter rows based on a corresponding condition to keep relevant information.
3. To rearrange and rename columns in the table based on a criteria.
4. To choose specific columns to keep in the table and create a new table with only the selected columns.

```python
q1_6 = ...
```

```python
grader.check("q1_6")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part 2: More Table Functions

<div class="alert alert-warning">
In Part 1 of this lab, we learned about several Table methods that help clean up and reorganize the appearance of a table to make our analysis easier. In this part, we'll learn about a couple more Table methods that help us filter and hone in on specific aspects of our data that we're interested in. In our case, these methods will help us better explore the breakdown of the gender and race/ethnicity demographics of students at UC Riverside.
</div>

### Sorting a column: `sort`

We will now introduce a new table method: `sort`. The `sort` table method allows us to see a table's column values sorted by its values in either **decreasing** (`descending=True`) or **increasing** (`descending=False`) order.

Let's say we want to ask the question: among all years, and particularly among the different groups of students based on their demographics, which fall headcount was the highest? Sorting the table will not only allow us to see a potential trend of headcounts over the years, but it will also show us if there are any trends among which groups of students have the highest headcounts. Let's use the `sort` method to find out:

```python
ucr_relabeled.sort("Fall Headcount", descending=True)
```

<div class="alert alert-warning">
The results that we see above are very interesting: firstly, we notice that the highest headcounts have all been in the most recent 10 years of the dataset (2013-2022), which shows us that student admissions at UC Riverside are higher than they were in the past; and secondly, we notice that the gender and race/ethnicity category is the same for the top 10 highest headcounts, female-identifying Chicanx/Latinx students. We may wonder how this breakdown would look by major or college as well, which we'll explore later in the lab.
</div>

### Question 2.1

Using the same method but changing the `descending` parameter to `False`, we can discover the lowest fall headcounts. Fill in the following code using the `sort` method to accomplish this.

*Note:* Be sure to use the `ucr_relabeled` table for this question!

```python
ucr_sorted_lowest = ...
ucr_sorted_lowest.show(5)
```

```python
grader.check("q2_1")
```

<div class="alert alert-warning">
The results from this call may not be as interesting as the previous one: we notice that there's a greater breakdown of full-time statuses and new/continuing statuses among these students, and each of these categories only had one individual in the 2022 fall semester. However, if we were to take a look at the next twenty or thirty lowest rows, this may yield some more valuable information about which groups of students are witnessing much lower admission rates.
</div>

---
### Row selection: `where` and the `are` Predicates

The last table method we will talk about is the `where` method. The `where` method keeps all rows that satisfiy a particular boolean condition. It takes in a column label and an `are` statement, which can be crafted using the `are` library. These are the most important `are` library predicates, but there are many more if you would like to investigate: [explore the 'are' library here.](https://data6.org/su25/reference/#tablewhere-predicates)

| Method | Input Type | Method Description |
| --- | --- | --- |
| `are.equal_to(n)` | number | Is the value from the column equal to `n`? |
| `are.above(n)` | number | Is the value from the column above `n`? |
| `are.above_or_equal_to(n)` | number | Is the value from the column above or equal to `n`? |
| `are.below(n)` | number | Is the value from the column below `n`? |
| `are.below_or_equal_to(n)` | number | Is the value from the column below or equal `n`? |
| `are.containing(s)` | string | Is `s` contained in the string value from the given column? |
| `are.containined_in(s)` | string | Is the string value from the given column contained in `s`? |

Adding a `not_` in front of all of these predicates makes each method do the opposite of what it does (ex: `are.not_equal_to(n)`).

*Note*: As we've seen in lecture, we can achieve an **exact match** by not explicitly using an `are` predicate. That is, `where("col", are.equal_to("something")` is identical to `where("col", "something")`; the latter is shorthand for the former.

For example, if we only wanted to see the rows in which the `"Fall Headcount"` was above 500, we could use the `where()` method combined with an `are` method:

```python
ucr_relabeled.where("Fall Headcount", are.above(500))
```

<div class="alert alert-warning">
In the context of our exploration of student demographics, this method can be incredibly useful. For example, if we wanted to delve further into fall headcounts that were above 3 (to avoid getting tons of random rows like we did when using the <code>sort</code> method) but below 20, in order to find groups of students that had relatively low values for fall headcounts, we could use the <code>are.strictly_between</code> predicate:
</div>

```python
ucr_relabeled.where("Fall Headcount", are.strictly_between(1, 21))
```

<div class="alert alert-warning">
From this view of our data, we might be able to see more interesting trends, such as the low values of full-time Native Hawaiian / Pacific Islander undergrads in 2022. We could also chain together multiple <code>where</code> calls in order to satisfy conditions for different rows in our table. If we wanted to get even more specific in looking at these lower headcounts, we could specify the values as we did previously, but then also look for only full-time students:
</div>

```python
ucr_relabeled.where("Fall Headcount", are.strictly_between(1, 21)).where("Full-Time Status", are.equal_to("Full-Time"))
```

<div class="alert alert-warning">
You may notice that some of the rows appear to be duplicates, but they are different -- we just can't see it because we dropped the <code>"New/Continuing Status"</code> column earlier. Now, when we really limit the data like this, we can see some trends about which races/ethnicities have less students belonging to them. This table still includes both transfers and frosh, undergraduates and graduates, but we can see some consistencies with Native Hawaiian/Pacific Islander, Domestic Unknown, and Black/African American students tending to have lower numbers. This type of more complex use of the <code>where</code> method showcases just how useful it can be.
</div>

In the example above, you can see the use of the predicate `are.equal_to` in order to match strings (you can also match exact numbers this way). You can also use a shorthand of omitting the `are` predicate completely to achieve the same effect. For example, if we wanted to see the rows of only undergrads, we could use the shorthand in this way:

```python
ucr_relabeled.where("Level", "Undergraduate")
```

When we omit the `are` predicate, the function is just using `are.equal_to` behind the scenes as the default option!

### Question 2.2
At this point in our analysis of the data, we have many tools under our belt to take a deeper look into the demographics of gender and race/ethnicity. For the following question, looking only at rows where the `"Year"` is equal to 2022, keep only the female-identifying students, and then sort by descending `"Fall Headcount"` value. Because this question is more complex and requires chaining together of multiple calls, we've provided skeleton code for you.

*Hint:* Use the shorthand in the example above with the `where` methods.

```python
ucr_f_2022 = ucr_relabeled.where(..., ...).where(..., ...).sort(..., ...)
ucr_f_2022.show(5)
```

```python
grader.check("q2_2")
```

<div class="alert alert-warning">
From the result above, you should see some similarities to what we found earlier when looking at the highest fall headcounts: female-identifying Chicanx/Latinx students make up some of the highest headcount values, regardless of <code>"New/Continuing Status"</code> or <code>"Original Entry Type"</code>. Additionally from this view, we can see that the next largest proportion of female-identifying students in 2022 had a race/ethnicity of Asian.
</div>

Now that we've learned a bit of surface-level information about the distribution of race/ethnicity, let's dig a little deeper into what the `"Gender"` column looks like. Utilizing the `np.unique` method that we used in the first section above, we can find the distinct values in the `"Gender"` column, to see which values are used to represent gender in this dataset.

```python
# Just run this cell
np.unique(ucr.column("Gender"))
```

As we can see above, the only two options for representing gender in this dataset are "Female" and "Male". It can often be hard to determine how to represent complex pieces of information (especially when this information is about individual people) as categorized versions in data.

<!-- BEGIN QUESTION -->

### Question 2.3 (Discussion)
Because this categorization of gender may not be the best way to represent this information, what do you think some potential issues with looking only at the female-identifying individuals from our data are? Additionally, what are the implications and potential sources of harm when we have to categorize information like this?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<div class="alert alert-warning">
We can tie your takeaways from this question back to one we asked earlier in the lab, while looking at the unique values available for race/ethnicity. There are certainly some issues with the representation of this demographic as well, as many individuals have to be lumped in together (e.g. students with origins from Middle Eastern or Arab countries having to identify themselves as "White", even when that likely does not accurately capture how they would identify themselves).
</div>

---

```python
from datascience import *
import numpy as np

import warnings
warnings.simplefilter('ignore')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part 3: Data Context Exploration

<div class="alert alert-warning">
As we saw at the end of Part 2, when working with data, it is always important to consider not only the impacts that the representation of the data can have in the real world</strong>, but also <strong>the effects that the conclusions from your data analysis and visualizations can have!</strong> While we were able to observe some interesting trends in Part 2, it is important to consider what types of conclusions we're making with our analysis and how applicable they are. In this part, we'll delve further into representations of data, while also utilizing additional data from UC Riverside broken down by college to try to analyze more interesting trends with student demographics.
</div>

### A Further Look Into the `Gender` Column

Let's continue with a bit more analysis on the `"Gender"` column. As stated above, the [source of the original data](https://ir.ucr.edu/stats/enrollment/demographic) from UC Riverside also included a second dataset that displays the same admissions data, broken down by college. An interesting question you may want to consider is why UC Riverside would make this type of breakdown available alongside the original dataset of the whole university. 

Let's load the dataset into a Table object below called `ucr_college`.

```python
ucr_college = Table.read_table("riverside_by_college.csv")
ucr_college.show(5)
```

As we can see from the `"Year"` column, this dataset also contains slightly more updated information, including data from 2023. Let's take a look at the way the data is represented in the `"Gender"` column for this dataset:

```python
np.unique(ucr_college.column("Gender"))
```

<!-- BEGIN QUESTION -->

### Question 3.1 (Discussion)
It looks like this dataset contains the additional values of `"Nonbinary"` and `"Unknown"` for `"Gender"`. Take a look at the "Details about the data" section [at the bottom of the website.](https://ir.ucr.edu/stats/enrollment/demographic) How was this particular variable measured? How might this impact our comparison between the conclusions we come to with this data vs. the data we previously worked with?

*Hint:* Think about what reducing the number of categories would mean.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Doing this type of extra basic background research on the source of our data can be incredibly helpful in our understanding of our data!

---
### Exploration of Aggregation by College

> "STEM Majors include those in the Bourns College of Engineering (BCOE) and the College of Natural and Agricultural Sciences (CNAS)."

Using this information, we can use the `where` method to create two different tables: one for STEM majors, and one for non-STEM majors. In the cell below, we generate the list of distinct college names at UC Riverside. As we see in the statement above, we can classify all non-STEM colleges as those in the list below, aside from the College of Engineering and the College of Natural and Agricultural Sciences.

```python
np.unique(ucr_college.column("College"))
```

### Question 3.2
In the two cells below, use the `where` method to create two different tables: one for STEM majors (`ucr_stem`), and one for non-STEM majors (`ucr_nonstem`), based on which college they belong to.

```python
ucr_stem = ...
ucr_stem.show(3)
```

```python
ucr_nonstem = ...
ucr_nonstem.show(3)
```

```python
grader.check("q3_2")
```

Now, if we wanted to do a rough tabular comparison of the STEM vs. non-STEM majors, we can group by `"Gender"` and compare the numbers.

For now, don't worry too much about understanding the syntax for the group method: we'll cover it in the next lab! The code in the two cells below simply aggregates the data by calculating the sums of the `"Fall Headcount"` column for each gender category.

```python
ucr_stem_grouped = ucr_stem.group("Gender", np.sum).select("Gender", "Fall Headcount sum").sort("Fall Headcount sum", descending=True)
ucr_stem_grouped 
```

```python
ucr_nonstem_grouped = ucr_nonstem.group("Gender", np.sum).select("Gender", "Fall Headcount sum").sort("Fall Headcount sum", descending=True)
ucr_nonstem_grouped 
```

<!-- BEGIN QUESTION -->

### Question 3.3 (Discussion)
From the above two tables, we can see an interesting discrepancy that we see quite often with this type of education-related data. What do you see when we compare these two tables?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

We can also visualize the data to be able to see the trends and proportions of the data more clearly. Below, we create a joined version of the headcounts split by gender for with one column containing the fall headcounts for STEM colleges, and another containing the headcounts for non-STEM colleges.

```python
ucr_joined_gender = ucr_stem_grouped.join('Gender', ucr_nonstem_grouped).relabeled(make_array('Fall Headcount sum', 'Fall Headcount sum_2'), make_array('STEM Fall Headcount', 'Non-STEM Fall Headcount'))
ucr_joined_gender.show(5)
```

<div class="alert alert-warning">We looked into some aspects of the students' race/ethnicities in Part 2, but now we can dig a little deeper by performing the same process of grouping our data and seeing the distribution between the STEM and non-STEM colleges. We start by creating the tables split by STEM and non-STEM, as well as the joined version with the headcounts of both categories. </div>

```python
ucr_stem_grouped_re = ucr_stem.group("IPEDS Race/Ethnicity", np.sum).select(["IPEDS Race/Ethnicity", "Fall Headcount sum"]).sort("Fall Headcount sum", descending=True)
ucr_stem_grouped_re 
```

```python
ucr_nonstem_grouped_re = ucr_nonstem.group("IPEDS Race/Ethnicity", np.sum).select("IPEDS Race/Ethnicity", "Fall Headcount sum").sort("Fall Headcount sum", descending=True)
ucr_nonstem_grouped_re 
```

```python
ucr_joined_re = ucr_stem_grouped_re.join('IPEDS Race/Ethnicity', ucr_nonstem_grouped_re).relabeled(make_array('Fall Headcount sum', 'Fall Headcount sum_2'), make_array('STEM Fall Headcount', 'Non-STEM Fall Headcount'))
ucr_joined_re.show()
```

<!-- BEGIN QUESTION -->

### Question 3.4
What do you notice when we compare these two groups of students, disaggregated by race?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<div class="alert alert-warning">
As we come to an end of our exploration of student demographics from UC Riverside, consider what we've learned about the different representations of categories, as well as what sorts of conclusions we were able to effectively draw from tabular vs. visualization methods. What more could we still do in order to more concretely explore student demographics? What type of information do you think is missing that could be useful? How might you change the way that some of these demographics are represented, if data collection could be more flexible and supportive of values that lie on a continuum or are subject to changes?
</div>

Going back to an idea we posed when starting to look into the `"Gender"` column: while this type of exploration is very important and can also be informative, it is also important to remember that it is not always enough to simply look at a trend like this and state it. Ultimately, it is typically very hard to encode data about humans into numbers and categories, **because in doing so, we lose information and context about the individual we are looking at.** In your analysis in the future, try to strike a balance of looking at interesting trends in the data and considering the original context of the data you're working with.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part 4: Table Methods with Movies

In this section, we'll be working with a `movies` data set that contains information about various American films over time. It contains the following columns:
1. `"Film"`: The name of the movie
2. `"Genre"`: The genre of the movie
3. `"Year"`: The year the movie was released
4. `"Lead Studio"`: The primary movie studio responsible for producing the movie
5. `"Audience score %"`: The score, out of 100%, given to the movie by viewers
6. `"Rotten Tomatoes %"`: The score, out of 100%, given to the movies by the website [Rotten Tomatoes](https://www.rottentomatoes.com/)
7. `"Worldwide Gross (Millions)"`: The total gross revenue, in millions of dollars, that the movie made
8. `"Quality"`: Descriptive ranking of the movie based on audience score

Let's load the dataset into a new table, `movies`, to get practice with the more advanced table methods.

```python
movies = Table.read_table("movies.csv")
movies.show(5)
```

## The [apply](http://data8.org/datascience/_autosummary/datascience.tables.Table.apply.html#datascience.tables.Table.apply) method

The `apply` method allows us to map a function's behavior onto an entire column of a table. We can use built-in Python functions (like `max`) or we can define our own functions and then *apply* those functions to the columns of a table.

The `apply` method takes at least 2 arguments. The first is a function, and the rest are as many column labels you need to run that function. The number of columns you need to specify is dependent on the number of arguments the function has. For example, if the function you provide needs two inputs, you need to list two columns for it to work on.

`apply` returns a NumPy array of the transformed values. We can ask questions like "How can I categorize the items in this column?" (like converting grade percentages into letter grades from lecture). We can also make modifications to a table, like rounding all the values in a column to a certain accuracy.

### Example Use
We can see some examples of `apply` at work. Let's use `apply` to take the average of the two score percentages, `"Audience score %"` and `"Rotten Tomatoes %"`.

```python
def average_score(audience_score, rt_score):
    "Computes the average score between the audience score and Rotten Tomatoes score"
    return (audience_score + rt_score) / 2
```

```python
average_scores = movies.apply(average_score, "Audience score %", "Rotten Tomatoes %")
average_scores
```

Now, let's add a new column called `"Average score %"` and populate it with the information we just assigned to `average_scores`. We'll re-assign this new table to `movies`.

```python
movies = movies.with_column("Average score %", average_scores)
movies.show(5)
```

We can use `apply` in this way to add columns to a table that transforms something in the table to something that more people can understand more easily!

### Input Type and Output Type

Very quickly, we should talk about **input type** and **output type**. **Input type** is the *type* of the arguments in a function. A function can take in as input any of the types we have talked about so far in this class (int, str, boolean, etc.). The input type is very important when using the table methods we talk about today, because there can be behavior we do not expect if we give a function the wrong input type. Let's see an example of this behavior when we try to use our `average_score` function on a row that has an **input type** of `str` instead of `int`:

```python
movies.apply(average_score, "Genre", "Lead Studio")
```

The error message does not make it *explicitly* clear that the function you provided got an incorrect input type, but if you are using a table method and the error message seems to indicate incompatible types interacting with each other, that may be the issue. You can see this on the **last line of the error message**:

```python
TypeError: unsupported operand type(s) for /: 'str' and 'int'
```

In simple terms, it means you're trying to *divide a string by an integer*, which you cannot do in Python.

### Question 4.1
Fill in the function `convert_to_dollars` which converts a dollar amount from *millions of dollars* to *dollars*. Then, use the `apply` method to convert all values from the `"Worldwide Gross (Millions)"` column into dollars. Finally, create a new column in the `movies` table called `"Worldwide Gross"` using the array resulting from your call to `apply`.

*Note*: The code for this question requires several steps. Feel free to create new cells to experiment!


```python
def convert_to_dollars(dollar_millions):
    "Converts a dollar amount from millions of dollars to dollars"
    ...

dollars = ...
movies = ...
movies
```

```python
grader.check("q4_1")
```

## The [group](http://data8.org/datascience/_autosummary/datascience.tables.Table.group.html#datascience.tables.Table.group) method

The `group` method is very helpful for organizing a table before asking more questions about it. you can think of the `group` method as organizing rows into bins based on one of their values. All the rows that share a column value go in the same bin! We saw this in action in Lab 3 when we grouped tables in order to visualize certain variables like `"Month"` or `"Binge Drink"` in the BRFSS data set. Now, we will learn how to use the `group` method ourselves.

`group` takes in 1-2 arguments. The first is a column label to group by, and the second is an optional function argument to group on, which defaults to counting the number of rows in the bin. We will see some examples of how this optional argument works below.

> For a visualization of the `.group` method, check out the [Data 8 Table Visualizer](http://www.data8.org/interactive_table_functions/)

Let's use the `group` method to organize our movies by `Genre` so that we can see which genres are in our table:

```python
movies_by_genre = movies.group("Genre")
movies_by_genre
```

### Question 4.2
What if we want to group by movie *quality*? Fill in the following cell with code that will assign `quality_groups` to a grouped table based on the `"Quality"` column.


```python
quality_groups = ...
quality_groups
```

```python
grader.check("q4_2")
```

We can use the optional `collect` argument to ask more complex questions instead of just asking for the number of rows in each bin. For example, to see the average `"Audience Score %"` of each bin, we can use `np.mean` or `np.average`.

*Note*: For our purposes, `np.mean` and `np.average` are exactly the same and you can use them interchangeably in the call to `group`. However, the function you choose will impact the **column names** of the resulting table.

```python
quality_group_ratings = movies.group("Quality", np.average)
quality_group_ratings
```

### Shortcomings of `collect` argument

Look at the `quality_group_rating` table above -- notice that we cannot specify which column we want the average of. As you can see, the `group` method will take the average of every column. For columns where it is **possible to take the average** (columns with the correct input type), it does so. For columns where taking the average doesn't work (with strings in this case), Python outputs nothing for those column averages.

Also, the column labels in this new table now have `average` at the end **except** the column label you grouped on. The name of the function you choose to group on will appear at the end of every label in the resulting table. **Like most table methods, the original table you are grouping does not change unless you reassign it**.

To get only the column we wanted the average of, we can use the `select` method to get the `Quality` and the `Audience score % average`:

```python
quality_ratings = quality_group_ratings.select("Quality", "Audience score % average")
quality_ratings
```

### Question 4.3
The table defaults to sort on the column you grouped on, which in this case is alphabetical order for strings. Oscar asks you to sort the `quality_ratings` table in decreasing order based on the average audience score. 

Write a line of code below that assigns `ratings_sorted` to a table that matches Oscar's request.


```python
ratings_sorted = ...
ratings_sorted
```

```python
grader.check("q4_3")
```

### Question 4.4 (Do Well-Liked Movies Make Money?) 
After hearing that you have a handy `movies` table, Rebecca asks you the following question:
>*Do well-liked movies make money?*

It's your job to answer Rebecca's question. To do so, create a new table called `money_by_quality` with the following **two columns**:
1. `"Quality"`: String describing the quality of a given movie
2. `"Worldwide Gross (Millions) mean"`: The average gross revenue for each movie quality

*Hint*: You may find the previous calls to `group` helpful.


```python
money_by_quality = ...
money_by_quality
```

```python
grader.check("q4_4")
```

Look at that! Using the `group` method, we found that **"Great"** movies make an average of 215 million dollars while **"Okay"** movies only make an average of 71.7 million dollars. Let's report back to Rebecca!

## The [pivot](http://data8.org/datascience/_autosummary/datascience.tables.Table.pivot.html#datascience.tables.Table.pivot) method

The `pivot` method allows us to see the *intersection* of two of our column labels. `pivot` essentially sorts the contents of the dataset based on the combination of the two column labels you pivot on. All the table's rows that share values in the pivoting columns all go into the same bin, and this happens for all combinations of the first and second column you are pivoting on.

The `pivot` method has 4 important arguments, 2 of which are mandatory and 2 of which are optional:

| **Argument** | **Description** |
| --- | --- |
| `columns` | The label whose unique values will appear as the **columns** of the outputted pivot table |
| `rows` | The label whose unique values will appear as the **rows** of the outputted pivot table |
| *Optional:* `values` | Values to use when aggregating |
| *Optional:* `collect` | Function used to aggregate the `values` provided in the previous argument |

You must use the two `values` and `collect` arguments together, one does not work without the other.

> For a visualization of the `.pivot` method, check out the [Data 8 Table Visualizer](http://www.data8.org/interactive_table_functions/)

The best way to understand how `pivot` works is to look at some examples and talk about what happens. Here we will pivot the `movies` table on its `Genre` and `Quality` column labels to see how many of each type of *genre* are in each category of *quality*:

```python
genre_quality_pivot = movies.pivot("Genre", "Quality")
genre_quality_pivot
```

The way we can read this table is almost how we read a graph in a math class. To see how many **Good Comedies** there are, we look at the row corresponding to `Good` and the column corresponding to `Comedy`, so there are 18 **Good Comedies**.

The default behavior of `pivot` is to just count the rows that appear at each intersection of the pivot. However, we can ask it to count another value in the table using any function we want! For example, if we want to know **the maximum amount of money** each intersection made instead of **how many** movies appear in each intersection, we can do that as well!

```python
genre_quality_average_grosses = movies.pivot("Genre", "Quality", values="Worldwide Gross (Millions)", collect=max)
genre_quality_average_grosses
```

Based on this small dataset, it seems that the highest grossing film is not actually a "Great" film but rather a "Good" one. Run the following cell to find out what film it actually is!

```python
movies.where("Worldwide Gross (Millions)", 709.82)
```

Ah, Twilight! Who can resist Robert Pattinson and Taylor Lautner?

### Question 4.5 (Does Studio Matter?) 
Now, let's write a query that can tell us what genre of movie each studio tends to make, and how much those movies made in total. We can use the `pivot` method just like we did above to see the breakdown of movies by studio and genre, and then aggregate those rows by adding up the worldwide gross of each intersection!

Assign `column_label` and `row_label` to column labels of `movies` that make the resulting `pivot` call have studios as the rows and genres as the columns.

Assign `value_to_collect` to a column label that we can use to collect the data we need to ultimately see the total amount of money made by each intersection of genre and studio. Use the column that uses millions.

Assign `collection_function` to a function you know that can add up all the movie earnings at each intersection


```python
column_label = ...
row_label = ...
value_to_collect = ...
collection_function = ...

studio_genre_total_gross = movies.pivot(column_label, row_label, value_to_collect, collection_function)
studio_genre_total_gross
```

```python
grader.check("q4_5")
```

## The [join](http://data8.org/datascience/_autosummary/datascience.tables.Table.join.html#datascience.tables.Table.join) method

The last method we will talk about in this lab is the `join` method. This method allows us to combine two different tables together!

The `join` method takes in 2 mandatory arguments and 1 optional argument:

| **Column Name** | **Description** |
| --- | --- |
| `column_label` | a column to use to join |
| `other` | another table |
| *Optional:* `other_label` | `other`'s label to join on (if not the same as `column_label`) |

If `other` has a label in common with the table you are joining with and this common label is the one you want to join on, then you do not need to use the optional argument. If you want to join on another column label or if neither table has a column label in common, then you can use the optional `other_label`.

The way join works takes some getting used to, so let's look at some examples of `join` at work!

We have the `dogs` and `owners` tables below, take a look at them a bit before we move on so you understand what data they contain:

```python
dogs = Table().with_columns(
    "Name", np.array(["Shefran", "Hero", "Fluffy", "Doge"]),
    "Breed", np.array(["Bichon Frise", "Shih-poo", "Corgi", "Coin"]),
    "Owner", np.array(["Su Min", "Su Min", "Josh", "Edwin"]),
)
dogs
```

```python
owners = Table().with_columns(
    "Owner", np.array(["Su Min", "Josh", "Edwin", "Sandra"]),
    "Owner Age", np.array([22, 21, 21, 20])
)
owners
```

As you can see, we have a column in common: `Owner`. Let's join these two tables together so that we can have all the doggy data in one place!

```python
doggy_data = dogs.join("Owner", owners)
doggy_data
```

This table now has all of our information in one place, which makes using it easier!

Now let's take a look at a more common example, where the column labels being named differently can cause a problem. We will use the exact same `dogs` and `owners` tables, but we will change a column label on `owners`:

```python
owners_new_label = owners.relabeled("Owner", "Name")

display(dogs, owners_new_label)
```

Now if we try to use the `join` method like we did last time, we run into an issue...

```python
doggy_data = dogs.join("Owner", owners_new_label)
doggy_data
```

Because the `owners` table does not have a column label called `Owner`, we may try to use the one column label they do have in common: `Name`...

```python
doggy_data = dogs.join("Name", owners_new_label)
doggy_data
```

...but this doesn't appear to work either. This is not an error, there is simply no table outputted by this join call!

This `join` call does not do what we want, because the `dogs` `Name` corresponds to the **dog's** name, but the `owners` `Name` corresponds to the **Owner's** name! No dog and owner have any of the same names, so there is no data to output in this `join` call!

Instead, we have to make sure we join on the `Owner` column from `dogs` and the `Name` column from `owners`! We can do this using the third *optional* argument in the `join` method:

```python
doggy_data = dogs.join("Owner", owners_new_label, "Name")
doggy_data
```

In this table, the `Name` column now refers to the name of the dog, and the `Owner` column corresponds to the name of the owner!

### Question 4.6 (A Slightly Different `Join`): 
The `join` method can also change the number of rows in its output. If there are multiple rows in one table that match with one row in the other, the `join` method will include rows for each of these matches in the output. Also, if there is a row in a table with no match in the other, there will be no row in the output that represents this row. Let's implement both these situations in practice and see how they work:

```python
# Just run this cell
# This new dogs table has a new extra dogs
more_dogs = dogs.with_rows([["Clifford", "Big Red", "Sandra"], ["Doug", "Golden Retriever", "Russell"]])
more_dogs
```

**Task**: Before we exectute the join between these tables, we should be able to calculate how many rows should there be in the output. Assign the variable `more_dog_owner_rows` to the number of rows that should result from joining `more_dogs` with `owners_new_label`. Run the cell below to see them again for clarity:

```python
display(more_dogs, owners_new_label)
```

```python
more_dog_owner_rows = ...
more_dog_owner_rows
```

```python
grader.check("q4_6")
```

### Confirm your answer

Remember, each row in the `more_dogs` table only gets a row in the output if it matches a row in `owners_new_label`

```python
# If an owner in the table has 2 dogs, both dogs should appear in the output
# If a dog has no owner in the owners table, the dog does not appear in the output
complex_doggy_data = more_dogs.join("Owner", owners_new_label, "Name")
complex_doggy_data
```

```python
# Doug is not in the new table, and both Sandra's dogs are present for a total of 5 rows
complex_doggy_data.num_rows
```

As we expect, both `Spot` and `Clifford` appear as Sandra's dogs, but `Doug` does not appear in the table (fear not, he still has an owner, but `Russell` is not in the owners table).

**Finally, for reference, here is the link to the Data 6 Python Reference (our Python cheat-sheet) so you can review some of the methods we've used for tables in this lab!**

[Python Reference](http://data6.org/su25/reference)

## Done! 😇

## Pets of Data 6
Lilo wishes you all the best!

<img src="lilo.jpeg" width="50%" alt="Two cute sleepy cats on a floral cushion"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1_2 (points: [0, 1])
```
>>> assert isinstance(ucr_rows, int)
```
```
>>> assert ucr_rows == 5021
```


Tests for q1_3 (points: [0, 1])
```
>>> assert isinstance(q1_3, np.ndarray)
```
```
>>> assert q1_3 == make_array(3)
```


Tests for q1_4 (points: [0, 0.5, 0.5, 1])
```
>>> assert isinstance(ucr_relabeled, Table)
```
```
>>> assert "Race/Ethnicity" in ucr_relabeled.labels
```
```
>>> assert "IPEDS Race/Ethnicity" not in ucr_relabeled.labels
```
```
>>> assert all(ucr_relabeled) == all(ucr.relabeled('IPEDS Race/Ethnicity', 'Race/Ethnicity'))
```


Tests for q1_5 (points: [0, 1, 1])
```
>>> assert isinstance(ucr_selected, Table)
```
```
>>> assert ucr_selected.num_columns == 5
```
```
>>> assert all(ucr_selected) == all(ucr_relabeled.select('Year', 'Level', 'Full-Time Status', 'Race/Ethnicity', 'Fall Headcount'))
```


Tests for q1_6 (points: [0, 1])
```
>>> assert isinstance(q1_6, np.ndarray)
```
```
>>> assert q1_6 == make_array(4)
```


Tests for q2_1 (points: [0, 1, 1, 1])
```
>>> assert isinstance(ucr_sorted_lowest, Table)
```
```
>>> assert ucr_sorted_lowest.take(0).column('Fall Headcount').item(0) == 1
```
```
>>> assert ucr_sorted_lowest.take(1).column('Race/Ethnicity').item(0) == 'Asian'
```
```
>>> assert all(ucr_sorted_lowest) == all(ucr_relabeled.sort("Fall Headcount", descending=False))
```


Tests for q2_2 (points: [0, 1, 1, 1])
```
>>> assert isinstance(ucr_f_2022, Table)
```
```
>>> assert ucr_f_2022.num_rows == 94
```
```
>>> assert all(ucr_f_2022.column('Gender') == 'Female')
```
```
>>> assert all(ucr_f_2022.column('Year') == 2022)
```


Tests for q3_2 (points: [1, 1, 1, 1, 1, 1])
```
>>> assert isinstance(ucr_stem, Table)
```
```
>>> assert isinstance(ucr_nonstem, Table)
```
```
>>> assert all(x in ["Bourns College of Engineering", "College of Natural and Agricultural Sciences"] for x in ucr_stem.column('College'))
```
```
>>> assert all(x in ["College of Humanities, Arts, and Social Sciences", "Graduate School of Education", "School of Business", "School of Medicine", "School of Public Policy"] for x in ucr_nonstem.column('College'))
```
```
>>> assert ucr_stem.num_rows == 7419
```
```
>>> assert ucr_nonstem.num_rows == 9451
```


Tests for q4_1 (points: 1)
```
>>> callable(convert_to_dollars)
True
```
```
>>> convert_to_dollars(10) == 10000000.0
True
```
```
>>> isinstance(movies, Table)
True
```
```
>>> 'Worldwide Gross' in movies.labels
True
```
```
>>> np.isclose(movies.column('Worldwide Gross').item(10), 376660000.0)
True
```


Tests for q4_2 (points: 1)
```
>>> isinstance(quality_groups, Table)
True
```
```
>>> quality_groups.num_rows == 3
True
```
```
>>> quality_groups.labels == ("Quality", "count")
True
```


Tests for q4_3 (points: 1)
```
>>> isinstance(ratings_sorted, Table)
True
```
```
>>> all(ratings_sorted.column("Quality") == make_array('Great', 'Good', 'Okay'))
True
```


Tests for q4_4 (points: 1)
```
>>> isinstance(money_by_quality, Table)
True
```
```
>>> money_by_quality.num_rows == 3
True
```
```
>>> money_by_quality.labels == ('Quality', 'Worldwide Gross (Millions) mean')
True
```
```
>>> np.isclose(money_by_quality.column(1).item(2), 71.72821428571429)
True
```


Tests for q4_5 (points: 1)
```
>>> isinstance(studio_genre_total_gross, Table)
True
```
```
>>> studio_genre_total_gross.num_rows == 12
True
```
```
>>> studio_genre_total_gross.labels == ('Lead Studio', 'Action', 'Animation', 'Comedy', 'Drama', 'Fantasy', 'Romance')
True
```
```
>>> studio_genre_total_gross.column("Comedy").item(5) == 59.19
True
```


Tests for q4_6 (points: 1)
```
>>> type(more_dog_owner_rows) in set([int, np.int32, np.int64])
True
```
```
>>> more_dog_owner_rows == 5
True
```
