```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 7 - Functions, Control, and Iteration

## Data 6, Summer 2025

Welcome to lab 7! This week we will be covering boolean operators, functions, **control**, **iteration**, and **strings**.

```python
# Just run this cell
from datascience import *
import numpy as np
import warnings
warnings.simplefilter('ignore')
```



# Boolean Operators

---

## Comparison Operators

Comparison operators are boolean operators that are used to compare two different values in Python.

We introduced six of these operators in lecture: `==`, `!=`,`<`, `<=`, `>`, and `>=`. Try using these operators on variables and values of all types, though in practice they are typically used with integers and floats, and sometimes on strings. Strings are treated in increasing alphabetical order: ('a' < 'b') == True.

Run the following series of cells to see some examples of these comparison operators in practice:

```python
1 < 2
```

```python
min(20, 30) <= max(-10, 20)
```

```python
(3 + 4) != (3 * 4)
```

```python
"abc" < "def"
```

### An Important Distinction

There is an important distinction here between `=` and `==`:

>`=` is used to **set** a Python name (i.e. assignment), while `==` is used as the **comparison operator** to compare two values. 

```python
# Set the name 'a' to have value 5
a = 5
```

```python
# Ask if it is equal to the value you think it should be
a == 5
```

## Membership Operators: `in`
There is one more boolean operator built into Python we have not mentioned yet. The keyword `in` allows you to check if strings are contained in larger strings. It can be used in other ways in Python, but for now we will use it to ask if the first string is a substring of the second string (if it appears in the other string).

Run the cells below for examples of `in`:

```python
"hello" in "hello world"
```

```python
"i" in "rhythm"
```

The `in` keyword can also check whether an array *contains* a certain element. `True` is returned if the array contains the element, and `False` otherwise.

```python
my_array = make_array(3,2,5,6)
5 in my_array
```

```python
9 in my_array
```

## Part 1: Boolean Operators: `and`, `or`, and `not`

`and`, `or`, and `not` are boolean operators that operate directly on booleans.

>`and` returns `True` when both operands are `True`, `False` otherwise.

>`or` returns `False` when both operands are `False`, `True` otherwise.

>`not` gives back the opposite boolean.

These three operators are useful when writing functions because they help control a function's logic. You will find them useful when writing more complicated functions in the course.

**Tip**: use parentheses to ensure your boolean operators are being applied correctly and without ambiguity.

Run the cells below to see these 3 operators in practice:

```python
(3 * 4 == 12) or (3 + 4 == 50)
```

```python
(3 * 4 == 12) and (3 + 4 == 50)
```

```python
not (3 + 4 == 50)
```

---
### Question 1

Write code that assigns `x_satisfies` to `True` or `False` depending on if the value of `x` satisfies *either* of the following conditions:
1. Even and has a value in the range 24 and 32 (inclusive)
2. Odd and has a value between 7 and 20 (exclusive). 

For this question, `x` will be assigned to `26`.

*Hint*: You should try using the modulo operator (%) to help check if a value is odd or even! The `%` operator will provide you with the *remainder* after diving by a certain number


```python
x = 26
x_satifies = ...
x_satifies
```

```python
grader.check("q1")
```

---
# Part 2: Functions

Functions allow us to combine the topics we have discussed so far in the course. Writing functions is useful when we want to use the same code repeatedly without wanting to write it out each and every time. You can think of functions like names -- they store lines of code instead of Python values, and they can be used when you need them.

Below is an example of a Python function that doubles its input:

```python
def double(x):
    return 2 * x
```

Let's break down the different parts of this function:
- The `def` at the beginning tells Python that we are making a new function.
- The blue `double` text is the name of our function. If we want to use this function, we will call it using this name.
- The `x` inside the parentheses is the input variable, otherwise known as a function's domain. A function can have as many input variables as you need, as long as they have different names.
- The `:` indicates we are done defining our function's name and input variables.
- All lines after this `def` line are **indented** to indicate we are is what we call the **body** of the function.
    - You can have as many lines as you want in your functions, but in this case we only have 1.
- The `return` keyword indicates we want to output the following evaluated Python expression. In our case, we want to multiply our input variable by 2.

Let's see what happens when we use our function:

```python
answer = double(5)
# We expect the value of this 'answer' variable to be 10 as long as our function is working correctly
answer
```

We can also call values that are assigned **outside** of the function, including other pre-existing functions or variables, **inside** a new function that we define.

However, any variable that we assign **inside** (also known as a "local" variable) cannot be called **outside** of the function, because it only exists in the function.

For example, using an outside variable inside a function:

```python
outside_value = "Hello, "

def greeting(place):
    return outside_value + place
```

```python
greeting("World!")
```

```python
def age_diff(birthday):
    """ This is called a docstring. This explains what the function does:
    Return the absolute value of the difference in ages from Cora and someone else."""
    cora_year = 2012
    return abs(birthday - cora_year)
```

```python
# Notice that, because cora_year exsits in the function, we cannot get it outside of the function
cora_year
```

---
### Question 2.1

Complete the function `triple_double`, so that it uses the function `double` we used earlier to double an input. Then, take that doubled input and triple it and return the triple-doubled value. 

For example: if we called `triple_double(2)`, the input 2 should be doubled to 4, and then tripled to 12. The function would then output 12.  


```python
def triple_double(x):
    ...
```

```python
grader.check("q2.1")
```

---
### Should I Stay or Should I Go
Now let's take a look at a function that uses **variables** and **boolean operators**:

Based on the color of the stoplight in front of us we want to know whether we should stop or keep going. Because we know we are going to encounter many lights on our trip, we should write a function! We can make two functions: one where we take risks and go on yellow lights, and one where we play it safe and stop for yellow lights

```python
def stop_or_go_risky(light):
    green = (light == 'green')
    yellow = (light == 'yellow')
    red = (light == 'red')
    return (green or yellow) and (not red) # Go on green and yellow, stop on red

def stop_or_go_safe(light):
    green = (light == 'green')
    yellow = (light == 'yellow')
    red = (light == 'red')
    return green and ((not yellow) and (not red)) # Go on green, stop on yellow and red
```

These functions are booleans because they output either True or False. We can write functions that output any type we need, and you will do so throughout the course!

Run the following cell as a demo:

```python
#We can use the outputs of these functions to let us know to stop the car if we need to on our way to the market.

light1 = stop_or_go_safe('green')
light2 = stop_or_go_safe('yellow')
light3 = stop_or_go_safe('red')

print("I am driving to the market! No rush!")
print("") # This is the same as hitting the return key while typing in Google Docs
print("Should I go at a green light?... " + str(light1))
print("Should I go at a yellow light?... " + str(light2))
print("Should I go at a red light?... " + str(light3))
print("")

#But if we find out we are running late, we may want to take more risks.
print("Oh no! I'm running late!")
print("")

light4 = stop_or_go_risky('green')
light5 = stop_or_go_risky('yellow')
light6 = stop_or_go_risky('red')

print("Should I go at a green light?... " + str(light4))
print("Should I go at a yellow light?... " + str(light5))
print("Should I go at a red light?... " + str(light6))
```

---
### Question 2.2: Hailstone

The Hailstone sequence from Hofstadter is as follows:
1. Pick a positive integer n as the start value.
2. If n is even, divide it by 2.
3. If n is odd, multiply it by 3 and add 1.
4. If you continue this process, n will eventually reach 1. 

Based on what you learned today, write a function that uses if/else statements to do **one step** of the hailstone sequence. It should take in an integer n and then return the corresponding Hailstone number for that value. For example, `hailstone(10)` should return 5 and `hailstone(9)` should return 28. 


```python
def hailstone(n):
    if ...:
        ...
    else:
```

```python
grader.check("q2.2")
```

```python
# Run this cell to see the Hailstone sequence in action!
print("Starting value: 10")
n1 = hailstone(10)
print("Step 1:", n1)
n2 = hailstone(n1)
print("Step 2:", n2)
n3 = hailstone(n2)
print("Step 3:", n3)
n4 = hailstone(n3)
print("Step 4:", n4)
n5 = hailstone(n4)
print("Step 5:", n5)
n6 = hailstone(n5)
print("Step 6:", n6)
```

# Part 3: Control

When writing functions in Python, we may want the function to behave differently depending on the input. We could choose to write several similar functions to accomplish this, but that would require copying and pasting much of the same code over and over again, only making small changes.

Instead, we need a way to tell one singular function to execute different code for different inputs. In lecture we discussed this idea as `if-else` statements.

Let's look at an example of a function that uses `if-else` statements to tell us what water looks like at different temperatures:

<img src='water.jpeg' width=300>

```python
def state_of_water(temperature):
    if temperature <= 32: # Water is solid at and below 32°F
        return "ice"
    elif temperature < 212: # Water is liquid between 32°F and 212°F
        return "liquid"
    else: # Water is gaseous above 212°F
        return "steam"
```

Here, we could have made 3 different functions that each dealt with a different state of water, but that would force us to know the state of water (what we currently do not know and are trying to figure out) so that we could pick the right function to use. This is why control such as `if-else` statements are so important!

Now let's use this function to determine the state of water in the following locations during the winter:

| **City** | Temperature (°F) |
| --- | --- |
| Berkeley | 58 |
| New York | 24 |
| Miami | 78 |
| Earth's Core | 10800 |

We can see the calculation each city's state of water using the function defined above:

```python
print(f"Water in Berkeley is: {state_of_water(58)}")
```

```python
print(f"Water in New York is: {state_of_water(24)}")
```

```python
print(f"Water in Miami is: {state_of_water(78)}")
```

```python
print(f"Water at the Center of the Earth is: {state_of_water(10800)}")
```

We can call the function on different inputs and the function deals with each input differently based on the control logic you gave it. This will be very important when you write functions that deal with large, sometimes unpredictable datasets where the logic of your function will deal with inputs you may not have directly prepared the function for.

# Part 4: Iteration

Iteration is the repetition of code. Typically, the variables in iteration change sequentially, and the iteration continues to repeat until a boolean condition indicates that the iteration be stopped. 

---
## Question 4

A `for` loop is an example of iteration in Python that we have seen in lecture. Let's use a `for` loop to write a function that finds the sum of the digits of a number.

As an example, the sum of the digits of the number 123 is 1 + 2 + 3 = 6, so `sum_digits(123)` should return 6.

**Some ideas you may want to keep in mind**:
- How can we 'peel off' a digit of a number? Consider using typecasting for this question. 
- Once we know how to do that, how can we do that repeatedly to all digits?


```python
def sum_digits(n):
    ...
```

```python
grader.check("q4")
```

# Part 5: Strings

We have already seen strings before in this course, but we want to introduce some new ideas with strings that will be very useful when dealing with data from real-world data sets.

Strings can be viewed as special lists of characters. They can be printed, sliced, indexed, and more! You can do all of these things with lists as well, so let's take a look at how strings and lists are similar:

```python
our_string = "Data Science at Berkeley rocks!"
print(our_string)
print(our_string[0] + our_string[5] + "@" + our_string[16:24])
```

We can iterate and slice strings too! In this way, strings are like special arrays.

There is an important fact to point out here: In Python,  `strings` are **immutable**. This means that you are not able to mutate (i.e. change) a string object in place. If you want to modify your string, you'll have to make an entirely new one.

```python
our_string
```

```python
# However you cannot do this with strings because they are immutable.
# This will be important when we talk about cleaning data sets before we can use them.
our_string[3] = 'x'
```

There are special string methods that help us edit strings, but we cannot forcefully change them with indexing or slicing. If you want to change a letter or a part of a string, you need to use those `str` methods from lecture.

---
## Question 5.1

Let's use this idea of array and string indexing and `if/else` to find all words that begin with a vowel. Complete the function `starts_with_vowels` so that it takes an array of strings, finds all of the words that begin with a vowel, and then returns an array of said words.

For the example array, `some_words`, your function should return `array(['alligator', 'aardvark', 'owl', 'unicorn'])`.

*Hint*: You'll need to use the `np.append` function, the documentation for which you can find [here]("https://numpy.org/doc/stable/reference/generated/numpy.append.html").

```python
some_words = make_array("alligator", "monkey", "zebra", "lion", "aardvark", "owl", "bear", "unicorn")

def starts_with_vowels(words):
    words_to_keep = make_array()
    for ... in ...:
        if ... in "aeiou":
            ...
starts_with_vowels(some_words)
```

```python
grader.check("q5.1")
```

---
## Question 5.2

Let's write a function that uses another `for` loop. We want to be able to find the acronym of a name. For example, we want to be able to call **The University of California, Berkeley** by an abbreviated `UCB`.

Because we don't care as much about small words even if they are capitalized, also make sure to **exclude words with three letters or fewer from your acronym**. That is why the `T` in `The` is not included in the example acronym above.

Write a function that takes in a string and returns its acronym.

_Hint_: the string method `.isupper()` will return `True` if the given string (or character) is entirely upper case, and `False` otherwise.

```python
# Here is a cell that explains the str.split() function used in the solution below:
sentence = "Hello, I am a student in Data 6!"
sentence.split()
```

```python
def acronym(title):
    # This str.split() function mentioned above turns a string into a list of its words.
    words = title.split()
    result = ""
    ...
    return result

acronym("The University of Califiornia, Berkeley")
```

```python
grader.check("q5.2")
```

## Done! 😇

## Pets of Data 6

**Obi** is a bit shy, but wishes you all the best!

<img src="obi.png" width="50%" alt="Two cute guinea pigs"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 1)
```
>>> isinstance(x_satifies, bool)
True
```
```
>>> x_satifies is True
True
```


Tests for q2.1 (points: 1)
```
>>> callable(triple_double)
True
```
```
>>> triple_double(2) == 12
True
```
```
>>> triple_double(3) == 18
True
```


Tests for q2.2 (points: 1)
```
>>> callable(hailstone)
True
```
```
>>> hailstone(2) == 1
True
```
```
>>> hailstone(3) == 10
True
```
```
>>> (hailstone(10) == 5) and (hailstone(9) == 28)
True
```


Tests for q4 (points: 1)
```
>>> callable(sum_digits)
True
```
```
>>> sum_digits(123) == 6
True
```
```
>>> sum_digits(0) == 0
True
```
```
>>> sum_digits(1234) == 10
True
```


Tests for q5.1 (points: 1)
```
>>> all(starts_with_vowels(some_words) == make_array('alligator', 'aardvark', 'owl', 'unicorn'))
True
```
```
>>> all(starts_with_vowels(make_array("bear")) == make_array())
True
```
```
>>> "zebra" not in starts_with_vowels(some_words)
True
```


Tests for q5.2 (points: 1)
```
>>> callable(acronym)
True
```
```
>>> acronym("The University of California, Berkeley") == "UCB"
True
```
```
>>> acronym("") == ""
True
```
```
>>> acronym("Data Science at Berkeley") == "DSB"
True
```
```
>>> acronym("The United States of America") == "USA"
True
```
