```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Homework 1 – Python Fundamentals

## Data 6, Summer 2025

In this homework assignment, you will apply what you've learned about arithmetic expressions, data types, boolean expressions, and functions. Your job is to write code anywhere you see `...`, as well as answer any open-ended questions.

This homework is due on **Tuesday July 15th, at 11:00 PM**. You must submit the assignment to Gradescope. Submission instructions can be found at the bottom of this notebook and on Ed. See the [syllabus](https://data6.org/su25/syllabus/#late-policy-and-extensions) for our late submission policy.

```python
# Just run this cell to import necessary Python packages
try:
    from datascience import *
except:
    %pip install -q datascience
    from datascience import *
import numpy as np
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Computation

<br/><br/>


<hr style="border: 1px solid #fdb515;" />

# Question 1 – Errors

In this question, you will identify why certain lines of code cause an error or have unexpected behavior.

---

## Question 1a

The following cell contains an error.

```python
8 = 10 - 2
```

Choose the best explanation of what's wrong with the code, and then assign the integer 1, 2, 3, or 4 to `reason_1a` below to indicate your answer.

1. Python already knows `8 = 10 - 2`.

2. `8` is a number, and numbers can't be Python names. In Python, "`x = 10 - 2`" means "assign `x` as the name for the value of `10 - 2`."

3. It should be `10 - 2 = 8`.

4. I don't get an error message. This is a trick question.

```python
# Your answer should be in the format reason_1a = n, where n is either 1, 2, 3, or 4

reason_1a = ...
```

```python
grader.check("q_1a")
```

---

## Question 1b

The following cell contains an error.

```python
four = 3
nine = four times four
```

Choose the best explanation of what's wrong with the code and assign the integer `1`, `2`, `3`, or `4` to `reason_1b` below to indicate your answer.

1. The `times` operation only applies to numbers, not the word `four`.

2. The name `four` cannot be assigned to the number 3.

3. four times four is sixteen, not nine.

4. The name `four` cannot be followed directly by another name.

```python
# Your answer should be in the format reason_1b = n, where n is either 1, 2, 3, or 4

reason_1b = ...
```

```python
grader.check("q_1b")
```

---

## Question 1c

Terry runs the following cell and expects that the value of `eight` will be `8`, stored as an `int`.

```python
eight = print(3 + 5)
```

What data type is the value bound to the name `eight`? In other words, does `eight` have type `int`, `float`, `str`, `bool`, or `NoneType`?

Assign your answer to the name `eight_type` below **as a string**. For example, if you think the answer is `int`, write `eight_type = 'int'` below.

```python
eight_type = ...
```

```python
grader.check("q_1c")
```

---

## Question 1d

Sam wants to create and print a string consisting of their name, age, and house address, as follows.

```python
name = 'Sam'
age = 51
house_number = 1436
house_street = 'Durant Avenue'

sam_greeting = 'My name is ' + name + ', I am ' + age + ' years old, and I live at ' + house_number + ' ' + house_street
print(sam_greeting)
```

However, it errors. In the cell below, assign `sam_greeting_fixed` to a fixed version of `sam_greeting` that doesn't error. Start by copying the right-hand side of the `sam_greeting` assignment statement above. Your answer must still involve `+`.

```python
sam_greeting_fixed = ...

# If your sam_greeting_fixed is correct, the following print statement should print
# 'My name is Sam, I am 51 years old, and I live at 1436 Durant Avenue'
sam_greeting_fixed
```

```python
grader.check("q_1d")
```

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

# Question 2 -  Names

We have seen how names can bind to values. They can also bind to function names! For example, consider the following cell, which binds `another_name_for_len` to the `len` function:

```python
# Demo: Just run this cell
another_name_for_len = len
s = "this string is length 24"
print(s)
print("length:", len(s))
another_name_for_len(s)
```

---

## Question 2a

In the below cell, assign `mystery_cast` to a function name such that when run, the last line `str(mystery_cast(5.00)) + '3'` evaluates to the string `'53'`.

**Hint**: Try casting!

```python
mystery_cast = ...
str(mystery_cast(5.00)) + '3'
```

```python
grader.check("q_2a")
```

---

## Question 2b

Assign `mystery_func` to a function such that when run, `g(mystery_func(2, g(h(1, 5), 3)), 4)` evaluates to the integer `10`.

```python
mystery_func = ...

# do not edit any code below this line
g = min
h = max
g(mystery_func(2, g(h(10, 5), 15)), 20)
```

```python
grader.check("q_2b")
```

<!-- BEGIN QUESTION -->

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

# Question 3 – Markdown

[Markdown](https://en.wikipedia.org/wiki/Markdown) isn't a programming language – it's a "markup language", which means that it is designed to format text to look nice. But like a programming language, Markdown has a set of rules (syntax) that govern how it works. Think of it like a fancy Google Docs.

Here are some of the most important rules in Markdown:

- Double asterisks (`**`) are how we make words bold. `My name is **Jordan**` renders to "My name is **Jordan**".
- Underscores (`_`) are how we make words italicized. `How are _you_?` renders to How are _you_?, and `what is **_up_**` renders to what is **_up_**.
- Hashtags (`#`, `##`, `###`) are how we create titles/headings. The more hashtags you include, the smaller the title is. The title for this question was generated using `## Question 5 – Markdown`.
- Bulleted lists are created using the (`-`) character.
- Numbered lists are created the same way they're created in a standard word processor, like Google Docs.

Markdown is important to learn because Jupyter Notebooks are often full of Markdown cells (just like this assignment). In future assignments, we will ask you to write answers to written questions using Markdown, and while the majority of your answers will be in plaintext (i.e. with none of the special characters above), learning Markdown is valuable nonetheless.

### Task

Your job is to recreate the output in the following screenshot, using the Markdown cell below. Feel free to replace "Bongo Burger" and "K-Pop" with your own favorite food and music genre, but make sure the formatting is the same.

<img src='markdown.png' width=300>

This is the first question in this homework where we will grade your work manually instead of relying on an autograder. There will be a few of these questions.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: The BRFSS Dataset

### Background on the Data

The dataset that we'll use in this next part of the homework comes from the Behavioral Risk Factor Surveillance System (BRFSS), a health survey fielded by the Centers for Disease Control and Prevention (CDC), which is the U.S. national public agency.

From the [BRFSS website](https://www.cdc.gov/brfss/index.html):
>The Behavioral Risk Factor Surveillance System (BRFSS) is the nation’s premier system of health-related telephone surveys that collect state data about U.S. residents regarding their health-related risk behaviors, chronic health conditions, and use of preventive services.

>By collecting behavioral health risk data at the state and local level, BRFSS has become a powerful tool for targeting and building health promotion activities. 

The dataset that you will investigate is a **subset of the 2020 BRFSS Survey**. We've taken all the data points corresponding to *fully-completed surveys* and, in our opinion, a few interesting columns. Since the entire data set is so large, we've randomly sampled 100,000 rows from the original data. While we've wrangled and cleaning the data set you'll use in your investigation, you're welcome to investigate the original source; you can do so via the [Survey Data section](https://www.cdc.gov/brfss/data_documentation/index.htm) of the BRFSS site.

<br/><br/>

### Loading Individual Data
 
 In this homework, we'll first start by looking at the BRFSS Survey Data on the **individual level**. That is, each row in our `brfss` table corresponds to a **unique individual** that responded to the survey. While these individuals are not identified by their government names, we can use a numerical indexing system to identify individuals in an anonymized way.

Run the following cell to load in the `brfss` table. It contains the following information (as columns):
1. `Index`: The unique identifier for an individual respondent
1. `Race/Ethnicity`: The race/ethnicity of the individual respondent
1. `Personal Doctor`: Whether or not the individual has someone they consider a personal doctor.
1. `Sleep Time`: The number of hours an individual sleeps, on average, in a 24-hour period
1. `Weight (lbs)`: The weight, in pounds, of the individual respondent
1. `Height (in)`: The height, in inches, of the individual respondent
1. `Data Weight`: The weight associated to the individual respondent to avoid biases in the survey process (to discuss in a later question)

We will explore each of these variables over the next few homework questions. To read about more variables as well as the survey questions themselves, please visit the [BRFSS 2020 Codebook](https://www.cdc.gov/brfss/annual_data/2020/pdf/codebook20_llcp-v2-508.pdf) and the [BRFSS 2020 Questionnaire](https://www.cdc.gov/brfss/questionnaires/pdf-ques/2020-BRFSS-Questionnaire-508.pdf).

```python
# Just run this cell to load the data
brfss = Table.read_table("brfss2020.csv")
brfss
```

<!-- BEGIN QUESTION -->

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

# Question 4 - Demographic Variables 

While the BRFSS measures many demographic variables (age, binary sex at birth, highest grade completed, etc.), for now let us focus on the `Race/Ethnicity` variable provided in the table above. The `Race/Ethnicity` variable in the BRFSS dataset takes on one of the following nine strings (i.e., categories):

* White, Non-Hispanic
* Black, non-Hispanic
* American Indian or Alaskan Native, Non-Hispanic
* Asian, Non-Hispanic
* Native Hawaiian or other Pacific Islander, Non-Hispanic
* Other race, non-Hispanic
* Multiracial, non-Hispanic
* Hispanic
* Missing/Don't Know

These categories are identical to those publicly released by the American Community Survey (ACS)!

In the cell below, write a few sentences that remark on this variable definition. Specifically, your comments should address:

1. Why do both the BRFSS (run by the US CDC) and the ACS (run by the US Census Bureau) define race/ethnicity using the same variable domain/categorization?
2. Why would a national health survey like the BRFSS measure race/ethnicity and other demographic data?

These readings may be useful: [About the BRFSS](https://www.cdc.gov/brfss/about/index.htm); p. 26 of the [BRFSS 2020 Questionnaire](https://www.cdc.gov/brfss/questionnaires/pdf-ques/2020-BRFSS-Questionnaire-508.pdf); 
 [About the ACS](https://www.census.gov/programs-surveys/acs/about.html); p. 2 of the [ACS 2020 Questionnaire](https://www2.census.gov/programs-surveys/acs/methodology/questionnaires/2020/quest20.pdf).

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

# Question 5 - Personal Doctor

In this question, you'll be investigating the following question from the BRFSS Survey:
> *Do you have one person you think of as your personal doctor or health care provider?*

Run the following cell to load in the `personal_doctor_array`, which contains the responses to the above question. The array is a numeric array containing values as `1` for "Yes" and `0` for "No" or "Missing Response." In other words, the "Personal Doctor" variable is a **binary variable.**

```python
# Just run this cell
personal_doctor_array = brfss.column("Personal Doctor")
personal_doctor_array
```

---

## Question 5a

Using the array `personal_doctor_array`, assign `has_personal_doctor` to `1` if the individual with index **37** (i.e., the 38th respondent) has a personal doctor, `0` if they don't/they did not respond.

**Note**: Do not "hardcode" the answer! Instead, you should use the array method `item(arg)` and supply the correct argument `arg` to index into the `personal_doctor_array` array.

```python
has_personal_doctor = ...
has_personal_doctor
```

```python
grader.check("q_5a")
```

---

## Question 5b

Using the array `personal_doctor_array`, assign `percent_personal_doc` to the **fraction** of people with a personal doctor in the dataset above. For the sake of this problem, an individual has a personal doctor only if they responded "yes" to the corresponding survey question.

You should **only** use Python built-in functions like `sum` and `len` for this problem! In other words, do **not** use any NumPy functions (e.g., do not use `np.count_nonzero`, etc.).

**Hint**: How can you make use of the *binary* nature of the "Personal Doctor" variable? The only values in the array are 1 (has personal doctor) 0 (has no known personal doctor).

```python
percent_personal_doc = ...
percent_personal_doc
```

```python
grader.check("q_5b")
```

<br/><br/>

<hr style="border: 1px solid #fdb515;" />


# Question 6 - Sleep

In this question, you'll be investigating another question from the BRFSS Survey:
> *On average, how many hours of sleep do you get in a 24-hour period?*

Run the following cell to load in the `sleep_col` array. This array contains the number of hours each respondent reported sleeping in a 24-hour period.

```python
# Just run this cell
sleep_col = brfss.column("Sleep Time")
sleep_col
```

It turns out that some values in `sleep_col` are `0`! Instead of printing out the entire array, run the below cell to print out a random set of 20 values from the array. You can run it multiple times to print out different sleep times.

```python
# Run this cell to print out 20 random sleep times
# Feel free to run it multiple times.
np.random.choice(sleep_col, 20)
```

Through the "random sampling" of sleep times above, you observe that some values in the `sleep_col` array are `0`. Since sleeping zero hours on average is near-impossible, this must be a **code** for something else. In this case, if the individual **did not respond** to this question, their hours of sleep value is `0`. (Note: In this context, "code" means a variable category, and not Python code).

### Task

Assign `sleeper_count` to an expression that uses a NumPy function to count the number of individuals who provided valid responses to the sleep survey question. Your expression should involve both a NumPy function and the array `sleep_col`.

**Hint**: Please see a list of NumPy functions in the [Data 6 Python reference](https://data6.org/su25/reference/). One of them will work!

```python
sleeper_count = ...
sleeper_count
```

```python
grader.check("q_6")
```

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

# Question 7 – BMI

Now, let's focus on the weight and height variables included in our BRFSS individual respondents table.

From [Wikipedia](https://en.wikipedia.org/wiki/Body_mass_index) the **Body Mass Index (BMI)** "is a convenient rule of thumb used to broadly categorize a person as underweight, normal weight, overweight, or obese based on tissue mass (muscle, fat, and bone) and height."

BMI is widely used in United States healthcare systems. Medical practitioners frequently consider a patient's BMI to determine eligibility for different medical services. We'll talk more about the social impacts of BMI later in this question, but first let's understand its computation:

BMI is defined according to the following formula:

$$\text{BMI} = \frac{\text{mass in kg}}{\text{height in meters}^2}$$

We'd like to compute the BMI for all individuals in our BRFSS dataset, but unfortunately the dataset only contains weights in pounds (lbs) and heights in inches (in). Here's the `brfss` table again for convenience.

```python
# just run this cell to show the table again
brfss.show(5)
```

 Additionally, we've provided two arrays, `weights` and `heights`, which contain the corresponding weights in pounds and heights in inches for all respondents.

```python
# Just run this cell -- you don't need to understand it yet!
non_zero_measurements = brfss.where("Weight (lbs)").where("Height (in)")

weights = non_zero_measurements.column("Weight (lbs)")
heights = non_zero_measurements.column("Height (in)")
print(weights)
print(heights)
```

Finally, we've provided the following conversion factors to convert from weight (lbs) to mass (kg), and from height (in) to height (m):

```python
# just run this cell to create conversion factors IN_TO_CM, CM_TO_M, LB_TO_KG
# Don't change any names/values in this cell!
IN_TO_CM = 2.54
CM_TO_M = 1 / 100
LB_TO_KG = 1 / 2.2
```

```python
# 1 lb is 0.4545 kg
LB_TO_KG
```

```python
# 1 in is 0.0254 m
IN_TO_CM * CM_TO_M
```

---

## Question 7a - Weight/Height Conversion (Individual)

Say we want to convert an individual's height and weight into the appropriate units for the above BMI calculation. For this calculation, let's use the individual corresponding to index **543** (i.e. the 544th respondent).

Using the above conversion factors, fill in the following code to determine this individual's height in meters (`indiv_height_m`) and mass in kilograms (`indiv_mass_kg`). Your answers should only involve names and arithmetic operators (`+`, `-`, `/`, etc.). No numerical constants.

```python
indiv_height_in = ...
indiv_weight_lb = ...

indiv_height_m = ...
indiv_mass_kg = ...

indiv_height_m, indiv_mass_kg
```

```python
grader.check("q_7a")
```

---

## Question 7b - BMI Computation (Individual)

Assign `indiv_bmi` to the BMI of the individual corresponding to index 543, calculated according to the formula provided at the beginning of this question.

```python
indiv_bmi = ...
indiv_bmi
```

```python
grader.check("q_7b")
```

<br/>

---

## Question 7c - BMI Computation (All Individuals)

In order to showcase the power of arrays, we can do more than just calculate the BMI of one individual. Rather, we can calculate the BMI of all individuals with one single calculation.

In the cell below, assign `bmi_array` to an array containing the calculated BMI for each individual within the `brfss` data set (and also assign `heights_m` and `masses_kg` accordingly). We've included the BMI formula again for your convenience:

$$\text{BMI} = \frac{\text{mass in kg}}{\text{(height in meters)}^2}$$

*Hint*: How can you translate what you implemented in Questions 7(a) and 7(b) to also work on arrays? You may want to use the `weights` and `heights` arrays (and remember their units)!

```python
heights_m = ...
masses_kg = ...

bmi_array = ...
bmi_array
```

```python
grader.check("q_7c")
```

---

## Question 7d - Context of BMI in US Society

The Body Mass Index (BMI) is widely integrated into the healthcare systems of the United States. However, its use is publicly debated: critics argue that the measure oversimplifies complex health and social factors; when used as a screening tool, it is an ineffective proxy for making healthcare decisions for individuals.

The following news article on BMI provides a particular historical and current context:

> [Article](https://drive.google.com/file/d/1zilh8q1oTjmK2M8uyNNoKZUFS0a7XtNA/view?usp=sharing). "Why BMI is a flawed health standard, especially for people of color." *Washington Post*, May 2021.

According to the article above, which of the below statements are correct takeaways? For each of the parts below, assign the provided name in the cell to the **string** `'correct'` or the **string** `'incorrect'`. For example, if Q2c(i)'s statement is a takeaway from the above article, `q2ci` should be assigned to `'correct'`; otherwise it should be assigned to `'incorrect'`.

### Question 7d(i)

BMI as a measure was originally established to measure the height and weight of an "average" White, European man.

```python
q7d_i = ...
```

```python
grader.check("q_7di")
```


### Question 7d(ii)

In the U.S. healthcare system, a patient's BMI is computed and checked against thresholds called "cut points," which are then used to determine eligibility for particular medications, surgeries, and even certain types of military service.

```python
q7d_ii = ...
```

```python
grader.check("q_7dii")
```


### Question 7d(iii)

Health institutions in the United States have universally adopted different cut points based on race and ethnicity, including for Latino people, Black people, and people of Asian descent.

```python
q7d_iii = ...
```

```python
grader.check("q_7diii")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 3: The BRFSS Dataset and the U.S. Population


Finally, let's explore how to aggregate the BRFSS data to different units of analysis. Before we do so, we need to introduce the statistical notion of **weighted averages**:

### Weighted Averages

Suppose we have five numbers: `0`, `2`, `5`, `10`, `12`.

The **arithmetic average** (i.e., mean) of a set of numbers is the sum of those numbers divided by the count of those numbers:

$$\text{average} = \frac{1}{5} \left( 0 + 2 + 5 + 10 + 12 \right)$$

Another perspective of the average is that each element contributes an *equal amount* to the final amount. In our example, each element contributes an equal 20%:

$$\text{average} = 0.2 (0) + 0.2 (5) + 0.2 (10) + 0.2 (12)$$

By contrast, a **weighted average** is when each element contributes a different amount (weight) to the final average, often based on its frequency of occurrence in a dataset. For example, if our elements respectively occurrred 20%, 30%, 10%, 0% and 40% in our data, then

$$\text{weighted average} = 0.2 (0) + 0.3 (2) + 0.1 (5) + 0 (10) + 0.4 (12)$$

where the weights are **frequencies** (and not percentages!) of occurrence.

### BRFSS Sleep Time Frequency Table

The below is a **frequency table** which aggregates individuals by their responses to the `"Sleep Time"` question from the BRFSS survey. While our previous `brfss` table had rows that represented **individual respondents**, each row in our frequency table represents a **group of individuals** with the same hours of sleep.

Run the following cell to load the `sleep` table. It contains the following columns:
1. `Sleep Time`: The number of hours slept in a 24-hour period, on average.
1. `# Respondents`: The number of survey respondents corresponding to the sleep time value
1. `% Respondents`: The % of survey respondents corresponding to the sleep time value
1. `% Population (Estimated)`: The **estimated population weights** corresponding to the each sleep time value (discussed below)

```python
# Just run this cell to load in the sleep arrays and % Population (Estimated) arrays
sleep = Table.read_table("sleep.csv")

sleep_times = sleep.column('Sleep Time')
percent_pops = sleep.column('% Population (Estimated)')

sleep.show(25)
```

### What Are Population Weights?

Note that other than `Sleep Time`, the other three variables are **aggregated** from the original data. The `# Respondents` and `% Respondents` are computed by counting the number of respondents who reported a valid daily sleep time (in hours) between 1 and 24 (or 0 if no survey response).

The `% Population` variable is the **estimated** percentage of population associated 
with the corresponding sleep time. These values are different from those of `% Respondents`! This is because the BRFSS 2020 collected data on about 300,000 individuals, which is much less than the 2020 US population of about 330,000,000 individuals.

The BRFSS computes **data weights** for each individual that estimates how common an individual's data are among the US population. These weights are likely computed through demographic data like age, gender, location, race/ethnicity, income, etc. The `Data Weight` variable from the original `brfss` table by individual was used to compute the `% Population` variable in this `sleep` table, aggregated by sleep time.

The `% Population` variable is useful in computing weighted averages, which we'll explore as our next (and last) homework question.

<br/><br/>

<hr style="border: 1px solid #fdb515;" />

# Question 8 - Weighted Sleep Average

We'd like to analyze sleep at the population level by computing the **weighted average sleep time** of the US population of 2020.

Instead of taking the simple average of sleep hours 0 through 24 (which would be 12 if we weighted all sleep times equally), we will weight sleep time by the estimated % population for that reported sleep time.

Assign `avg_sleep` to the weighted average daily sleep time of the US population. In other words, you should compute a weighted average of the `Sleep Time` and `% Population (Estimated)` columns.

**Hint**: The average sleep time should be between 0 and 24 (the number of hours in a day). Remember to use frequencies and not percentages!

```python
# We've provided these arrays for you
sleep_array = sleep.column("Sleep Time")                  # Don't change this line
weights_array = sleep.column("% Population (Estimated)")  # Don't change this line

avg_sleep = ...
avg_sleep
```

```python
grader.check("q_8")
```

**(Bonus, ungraded)** The weighted average we computed above actually *underestimates* the average sleep time You may have noticed that the weighted average you computed above includes 0 sleep hours, which, as we previously discussed, is a code for "no response." How would you recompute the weights to only take the average of *positive* sleep times? Use the cell below!

```python
# Use this cell to answer the bonus question!
```

## Done!

Congratulations! You've finished your first Data 6 homework assignment! To submit your work, follow the steps outlined on Ed.

| Category | Points |
| --- | --- |
| Autograder (Coding questions) | 20 |
| Written (Q3 and Q4) | 4 |
| **Total** | 24 |

## Pets of Data 6
Take it easy like Tigger. Congrats on finishing the first Data 6 homework!

<img src="tigger.jpg" width="50%" alt="Two cute sleepy cats on a floral cushion"/>

## Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by colicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. Either way, **you are responsible for ensuring your submision follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'hw01'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True)
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q_1a (points: [0, 0])
```
>>> isinstance(reason_1a, int)
True
```
```
>>> reason_1a in [1, 2, 3, 4]
True
```


Tests for q_1b (points: [0, 0])
```
>>> isinstance(reason_1b, int)
True
```
```
>>> reason_1b in [1, 2, 3, 4]
True
```


Tests for q_1c (points: [0, 0])
```
>>> isinstance(eight_type, str)
True
```
```
>>> eight_type in ['int', 'float', 'str', 'bool', 'NoneType']
True
```


Tests for q_1d (points: [1])
```
>>> isinstance(sam_greeting_fixed, str)
True
```


Tests for q_2a (points: [0.5, 0.5])
```
>>> isinstance(mystery_cast, type)
True
```
```
>>> str(mystery_cast(5.0)) + '3' == '53'
True
```


Tests for q_2b (points: [1])
```
>>> g(mystery_func(2, g(h(10, 5), 15)), 20) == 10
True
```


Tests for q_5a (points: [0.5])
```
>>> has_personal_doctor in make_array(0, 1)
True
```


Tests for q_5b (points: [1, 1])
```
>>> isinstance(percent_personal_doc, float)
True
```
```
>>> np.isclose(percent_personal_doc, 0.83341)
True
```


Tests for q_6 (points: [1, 1])
```
>>> isinstance(sleeper_count, int)
True
```
```
>>> sleeper_count == 98907
True
```


Tests for q_7a (points: [0.25, 0.25])
```
>>> isinstance(indiv_height_m, float)
True
```
```
>>> isinstance(indiv_mass_kg, float)
True
```


Tests for q_7b (points: [0])
```
>>> isinstance(indiv_bmi, float)
True
```


Tests for q_7c (points: [0])
```
>>> isinstance(bmi_array, np.ndarray)
True
```


Tests for q_7di (points: [0])
```
>>> isinstance(q7d_i, str)
True
```


Tests for q_7dii (points: [0])
```
>>> isinstance(q7d_ii, str)
True
```


Tests for q_7diii (points: [0])
```
>>> isinstance(q7d_iii, str)
True
```


Tests for q_8 (points: [1])
```
>>> isinstance(avg_sleep, float)
True
```
