```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 3 – Functions Practice

## Data 6, Summer 2025

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

```python
# Import necessary packages
from datascience import *
import numpy as np
```

<hr style="border: 1px solid #fdb515;" />

The table `survey_responses` contains responses from students along with the amount of time (in minutes) they took to complete the survey.

```python
survey_responses = Table.read_table("survey_responses.csv")
survey_responses
```

<br/>

---

## Question 1: Fun with Functions!

Let's do some survey analysis! Implement a function that takes in a Table `table` and an integer `threshold`, and returns a Table with all the survey responses that spent longer or equal to `lower_threshold` minutes and less than `upper_threshold` to fill out the survey.

```python
def filter_time(table, lower_threshold, upper_threshold):
    """
    Takes in a table and returns all the rows that spent longer than threshold minutes.

    Inputs:
    table: a Table that contains a column 'Time Taken (minutes)'
    threshold: an integer

    Returns:
    a Table with all the survey responses that spent longer than threshold minutes
    """

    ...
    
```

```python
grader.check("q1")
```

```python
between_ten_and_twenty = filter_time(survey_responses, 10, 20)
between_ten_and_twenty
```

<br/>

---

## Question 2: String Manipulation

Jedi wants one string of all of the `Student Name`s. Write a function that takes in a table and a column name, and concatenates together all of the names, separated by a parameter `separator`.


```python
def concatenate_names(tbl, separator, col_name):
    """
    Inputs:
    tbl: a Table
    separator: a string
    col_name: a string

    Returns:
    a string with all the names concatenated together

    """
    ...

# Now call the function with the survey_responses table, a space as the separator, and "Student Name" as the column name
concatenated_survey_names = ...
concatenated_survey_names
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

**Reflect:** What information did we lose when concatenating all the survey names together into one string?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Question 3: Functions within functions.

Sometimes, we'll have helper functions to perform some subroutine for us so we can abstract away more low-level tasks. Finish this implementation of a function that adds a new column, `Updated Values`, to a table that doubles, then squares all the values inside of column `col`.

```python
def add_doubled_squared_col(tbl, col):
    """Adds a new column called Updated_Values to tbl that doubles, then squares the values in col"""
    ...


survey_with_modified_times = add_doubled_squared_col(survey_responses, "Time Taken (minutes)")
survey_with_modified_times
```

```python
grader.check("q3")
```

Now, here is the re-implementation of the helper function without taking in any parameters.

```python
def add_doubled_squared_col(tbl, col):
    """Adds a new column called Updated_Values to tbl that doubles, then squares the values in col"""
    def double_squared():
        """Takes in nothing and returns a new array with the values doubled, then squared"""
        return (tbl.column(col) * 2) ** 2
    return tbl.with_column("Updated Values", double_squared())


survey_responses.column("Student Name")
```

```python
survey_with_modified_times = add_doubled_squared_col(survey_responses, "Time Taken (minutes)")
survey_with_modified_times
```

<!-- BEGIN QUESTION -->

---

## Question 4: Understanding Error Messages

For each of the following buggy implementations, write what the issue is and fix the implementation.

```python
def triple(x):
    """Return the value of x, tripled"""
    x * 3
```

```python
a = 5
triple(a) #We expect this to evaluate to 15
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

```python
def quadruple(x):
    """Return the value of x, quadrupled"""
    4 = x * 4
    return 4

quadruple(4)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

```python
def quintuple(x):
    x = x * 5
    return x

quintuple(4)
print(x) # Jedi expects x to be 20
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr>

## Done! 😇

## Pets of Data 6
Enjoy the grass and shade during this Summer heat! Congrats on finishing Lab 3.

<img src="paulina.jpg" width="50%" alt="Two cute sleepy cats on a floral cushion"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 1)
```
>>> sum(filter_time(survey_responses, 10, 20).column("Time Taken (minutes)")) == 82
True
```
```
>>> filter_time(survey_responses, 10, 20).num_rows == 6
True
```
```
>>> sum(filter_time(survey_responses, 90, 110).column("Time Taken (minutes)")) == 100
True
```


Tests for q2 (points: 1)
```
>>> concatenate_names(survey_responses, " ", "Student Name") == 'Oski Bruin Triton LoFi Hip Hop Girl Jedi Aurora Phoenix Cascade River Aspen Comet Galaxy Nebula Echo Sierra Horizon Summit Zephyr Blaze Indigo'
True
```
```
>>> concatenate_names(survey_responses.take(0), " ", "Student Name") == 'Oski'
True
```


Tests for q3 (points: 1)
```
>>> "Updated Values" in survey_with_modified_times.labels
True
```
```
>>> survey_with_modified_times.column("Updated Values").sum() == 51278
True
```
```
>>> add_doubled_squared_col(survey_responses.take(0), "Time Taken (minutes)").column("Updated Values").item(0) == 100
True
```
