```python
# Initialize Otter
import otter
grader = otter.Notebook("disc06.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Gini Coefficient

Now, let's dive deeper into the **Gini coefficient**. The Gini coefficient is a number between **0 and 1**, where **0** signals perfect equality between groups and **1** signals perfect inequality. The gini coefficient is used to measure the gender inequality between the gender groups, in this case men and women. If we assume the population is 50% male and 50% female, then gender inequality is when each gender deviates from their 0.5 share. Therefore, gender inequality = $|s_F - s_M|$, where $s_F$ (female share) is the number of women with characteristic $Y$ divided by the total population with $Y$, and $s_M$ (male share) is the number of men with $Y$ divided by the total population with $Y$ — assuming the population of genders was equal.

However, since this is typically not the case (there are not equal numbers of men and women), to find the Gini coefficient for characteristic $Y$, we use the formula:

$$
\sum_g |s_g - p_g|
$$

This equation sums over all genders, which in this case are males and females. It takes $s_g$ (gender $g$'s share), where $s_g$ is the number of people in gender $g$ who have characteristic $Y$ divided by the total number of people with $Y$, and subtracts $p_g$ (gender $g$'s population share), where $p_g$ is the number of people in gender $g$ divided by the total population across all genders.

Since shares and populations each make up a whole, they must sum to 1. Therefore, $s_M$ and $s_F$ are complements of each other (i.e. $s_F = 1 - s_M$), and the same applies to $p_F$ and $p_M$.


```python
import pandas as pd
from datascience import *
import matplotlib.pyplot as plt
%matplotlib inline
import seaborn as sns
import numpy as np
```

<hr style="border: 1px solid #fdb515;" />

## Gini Coefficient Dataset

```python
adult_gini = Table().with_columns(
    "Year", [1970, 1975, 1980, 1985, 1990, 1995, 2000],
    "share_f", [0.448, 0.449, 0.454, 0.460, 0.462, 0.465, 0.477],
    "pop_m", [0.496, 0.498, 0.497, 0.496, 0.498, 0.497, 0.495]
)
adult_gini
```

### Question 1

The `adult_gini` tables shows the female share (`share_f`) of literate adults and male population (`pop_m`). To be able to calculate the adult literacy gini coefficient for each year, we need to know the male share of literate adults and female population. Using the information above, calculate `share_m` and `pop_f`. Assume the table is assigned to `adult_gini`.

**Note:** It's important to know that `share_f` + `share_m` = 1 and `pop_f` + `pop_m` = 1. 

```python
pop_f = ...
share_m = ...
adult_gini = ...
adult_gini
```

```python
grader.check("q1")
```

### Question 2

Now that we have all the necessary information, use the formula to calculate **adult_coeff**, the gini coefficients for each year. Add the column `gini_coefficient` to the `adult_gini` table.

```python
adult_coeff = ...
adult_gini = adult_gini.with_columns("gini_coefficient", adult_coeff) 
adult_gini
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

### Question 3 (Discussion)

This equation sums over all genders (men and women). What kind of implications does this entail?

_Type your answer here, replacing this text._

```python
grader.check("q2")
```

<!-- END QUESTION -->

## Creating Visualizations

Using `gini_coefficients`, write the corresponding line of code that will produce the following graphs.

```python
gini_coefficients  = Table.read_table("student_ginis.csv") 
gini_coefficients 
```

<!-- BEGIN QUESTION -->

### Question 4

![Plot 1](image.png)

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 5

![Plot 2](image2.png)

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 6

![Plot 3](image3.png)

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 7

![Plot 4](image4.png)

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 8
![Plot 5](image5.png)

```python
...
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 1)
```
>>> pop_f == 1 - adult_gini.column("pop_m")
array([ True,  True,  True,  True,  True,  True,  True], dtype=bool)
```
```
>>> share_m == 1 - adult_gini.column("share_f")
array([ True,  True,  True,  True,  True,  True,  True], dtype=bool)
```
```
>>> "pop_f" in adult_gini.labels
True
```
```
>>> "share_m" in adult_gini.labels
True
```


Tests for q2 (points: 1)
```
>>> "gini_coefficient" in adult_gini.labels
True
```
```
>>> expected = make_array(0.112, 0.106, 0.098, 0.088, 0.08, 0.076, 0.056)
>>> np.allclose(adult_gini.column("gini_coefficient"), expected, rtol=1e-4)
True
```
