```python
# Initialize Otter
import otter
grader = otter.Notebook("tut04.ipynb")
```

<img src="data6.png" alt="Data Science Mentors Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Tutoring – Dictionaries

## Data 6

In this week's tutoring we will be going over dictionaries! 

Resources: 
* [Python Reference Sheet](https://data6.org/su25/reference/) 
* [Python Tutor](https://pythontutor.com/python-compiler.html#mode=edit)

---

### Dictionary Quick Guide:

We first define the following dictionary:

```python
shopping_list = {'apples': 4, 'bananas': 6, 'bread': 1}
```

Then we can use the following dictionary methods:

| Method                  | Description                                      | Example Output                                                |
|-------------------------|--------------------------------------------------|---------------------------------------------------------------|
| `shopping_list.keys()`  | Returns all **keys** in the shopping list.          | `dict_keys(['apples', 'bananas', 'bread'])`                |
| `shopping_list.values()`| Returns all **values** in the shopping list.     | `dict_values([4, 6, 1])`                                      |
| `shopping_list.items()` | Returns **key-value** pairs as tuples.           | `dict_items([('apples', 4), ('bananas', 6), ('bread', 1)])`   |

**Note:** You can iterate through all three different outputs above.

```python
# Imports
from datascience import *
import numpy as np
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part 1: 🍝 Italian Brainrot: Dictionary Practice

The Italian Brainrot meme trend features absurd, AI-generated characters with over-the-top names and designs. In this section, we’ll use a dictionary to represent and explore these characters, practicing key methods like `.keys()`, `.values()`, and `.items()`. 

Below, we have a nested dictionary where each character name is a key, and the value is another dictionary containing details like their species, signature item, and how many videos they appear in.

<hr style="border: 1px solid #fdb515;" />

```python
brainrot = {
    'Tralalero Tralala': {
        'species': 'shark',
        'item': 'nike sneakers',
        'videos': 12
    },
    'Bombardiro Crocodilo': {
        'species': 'crocodile airplane',
        'item': 'breadstick launcher',
        'videos': 18
    },
    'Ballerina Cappuccina': {
        'species': 'ballerina with a coffee cup head',
        'item': 'espresso twirler',
        'videos': 7
    }
}
```

<hr style="border: 1px solid #fdb515;" />

## Question 1a: Accessing Information From a Nested Dictionary

What is the `species` of `Tralalero Tralala`? Use code to access the value from the `brainrot` dictionary.

```python
tt_species = ...
tt_species
```

```python
grader.check("q1a")
```

<hr style="border: 1px solid #fdb515;" />

## Question 1b: Adding a Character

Add a new brainrot character named `Tun Tun Sahur` to the dictionary. This character is a **dancing toaster** with a **ladle**, and appears in **9** videos. Use dictionary syntax to add this new entry to the existing `brainrot` dictionary.

```python
brainrot["Tun Tun Sahur"] = ...
brainrot
```

```python
grader.check("q1b")
```

<hr style="border: 1px solid #fdb515;" />

## Question 1c: Counting

Write code to compute the total number of videos all characters have appeared in and assign the number to `total`.

**_Hint:_** A loop may be useful here

```python
total = 0
...
    ...
total
```

```python
grader.check("q1c")
```

## Question 1d: `.items()`

Make a list of all characters with more than 10 videos
Use .items() and a loop to create a list called popular_chars that includes only the names of characters who appear in more than 10 videos.

```python
popular_chars = []
for name, info in brainrot.items():
    ...
        ...
popular_chars
```

```python
grader.check("q1d")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1a (points: 1)
```
>>> tt_species == 'shark'
True
```


Tests for q1b (points: 1)
```
>>> brainrot.get("Tun Tun Sahur") == {"species": "dancing toaster", "item": "ladle", "videos": 9}
True
```


Tests for q1c (points: 1)
```
>>> total == 46
True
```


Tests for q1d (points: 1)
```
>>> 'Tralalero Tralala' in popular_chars
True
```
```
>>> 'Bombardiro Crocodilo' in popular_chars
True
```
```
>>> len(popular_chars) == 2
True
```
